---
# Welcome to CE880
### This is your week-7 : Problem notebook

For this problem set, we'll be using the Jupyter notebook and please upload this notebook to [Google Colab](https://colab.research.google.com/). 

---
## Question 1: 

#### Bootstrapping

We are generating a random vector, which will be our population, whose law of distribution (expected to be the same as that of the population) is unknown. Write a function to pick a sample from our population and by applying the Bootstrapping procedure, where sample `size=300`: and `B=1000` number of re-sampling

<!-- grade_id: my_bootstrapping_ans, points: 0, grade: false, solution: true -->
```python
import numpy as np
np.random.seed(123)
pop = np.random.randint(0,500 , size=100)
print(pop.mean()) # Print population mean
sample = np.random.choice(pop, size=300) #so n=300
def my_bootstrapping():
    """Write a function to pick a sample from our population and by applying the
    Bootstrapping procedure, where sample size=300: and B=1000 number of re-sampling """
    # YOUR CODE HERE
    sample_mean = []
    for i in range(1000):
      sample1 = np.random.choice(sample, size=300)
      avg = np.mean(sample1)
      sample_mean.append(avg)
    print(np.mean(sample_mean))
    return np.mean(sample_mean)
    raise NotImplementedError()
```

Output:
```
253.76
```

<!-- grade_id: my_bootstrapping_test, points: 2, grade: true, solution: false -->
```python
# Check you solution by running this cell
import math
assert math.isclose(my_bootstrapping(), 253.241, rel_tol = 0.05)
```

Output:
```
252.79019
```

---
## Question 2: 

#### Majority Voting
Given a list = `[1,2,3,1,2,1,1,1,3,2,2,1]`. Write a function to find the majority voting.

```python
import collections
a = np.array([1,2,3,1,2,1,1,1,3,2,2,1])

def my_majorityvote(a):
    """Write a function to find the majority voting."""
    # YOUR CODE HERE
    counter = collections.Counter(a)
    majority = max(counter, key=counter.get)
    return majority
```

<!-- grade_id: my_pairedttest_ans, points: 0, grade: false, solution: true -->
```python
# Example of the Paired Student's t-test
a = np.array([1,2,3,1,2,1,1,1,3,2,2,1])

def my_majorityvote():
    """Write a function to find the majority voting."""
    # YOUR CODE HERE
    majority = None
    majority_count = 0
    for vote in a:
        vote_count = 0
        for other_vote in a:
            if vote == other_vote:
                vote_count += 1
        if vote_count > majority_count:
            majority = vote
            majority_count = vote_count
    print(majority)
    return majority
    raise NotImplementedError()
```

<!-- grade_id: my_pairedttest_test, points: 2, grade: true, solution: false -->
```python
# Check you solution by running this cell
import math
assert math.isclose(my_majorityvote(), 1, rel_tol = 0.05)
```

Output:
```
1
```

---
## Question 3: 

#### Decision Tree
USe the IRIS dataset from `sklearn` to build a decision tree classifier. The dataset contains information for three classes of the IRIS plant, namely IRIS Setosa, IRIS Versicolour, and IRIS Virginica, with the following attributes: sepal length, sepal width, petal length, and petal width. Your aim is to train a decision tree on training data and return the test accuracy.

<!-- grade_id: my_DTree_ans, points: 0, grade: false, solution: true -->
```python
import pandas as pd
import numpy as np
from sklearn.datasets import load_iris
from sklearn.metrics import accuracy_score
from sklearn.tree import DecisionTreeClassifier

# Reading the Iris.csv file
data = load_iris()

# Extracting Attributes / Features
X = data.data

# Extracting Target / Class Labels
y = data.target

# Import Library for splitting data
from sklearn.model_selection import train_test_split

# Creating Train and Test datasets
X_train, X_test, y_train, y_test = train_test_split(X,y, random_state = 50, test_size = 0.25)

def my_DTree():
    """Write code to to train a decision tree classifier and return the test accuracy"""
    # YOUR CODE HERE
    # instantiate the DecisionTreeClassifier model with criterion gini index
    clf_gini = DecisionTreeClassifier(criterion='gini', max_depth=3, random_state=0)
    # fit the model
    clf_gini.fit(X_train, y_train)
    y_pred_gini = clf_gini.predict(X_test)
    score = accuracy_score(y_test, y_pred_gini)
    print('Model accuracy score with criterion gini index: {0:0.4f}'. format(accuracy_score(y_test, y_pred_gini)))
    return score
    raise NotImplementedError()
```

<!-- grade_id: my_DTree_test, points: 2, grade: true, solution: false -->
```python
# Check you solution by running this cell
import math
assert math.isclose(my_DTree(), 0.94736, rel_tol = 0.05)
```

Output:
```
Model accuracy score with criterion gini index: 0.9474
```

```python

```